![Final Lesson Exercise](images/Banner_FEX.png)

# Lesson #9: Cross validation and classification
## Edible and Poisonous Mushrooms dataset

## About this assignment
In this assignment, you will explore information regarding mushrooms.<br/>

This time you will practice the cross validation techniques to select best classifiers.<br />

## Preceding Step - import modules (packages)
This step is necessary in order to use external packages. 

**Use the following libraries for the assignment, when needed**:

```python
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# === CELL TYPE: IMPORTS AND SETUP 

import os                       # for testing use only

import numpy as np
import pandas as pd
from sklearn.model_selection import train_test_split
# --------cross-validation
from sklearn.model_selection import cross_val_score
from sklearn.model_selection import GridSearchCV
# -------- classification
import sklearn
from sklearn import neighbors, tree, ensemble, naive_bayes, svm
# *** KNN
from sklearn.neighbors import KNeighborsClassifier
# *** Decision Tree; Random Forest
from sklearn import tree
from sklearn.tree import DecisionTreeClassifier
from sklearn.ensemble import RandomForestClassifier
# *** Naive Bayes
from sklearn.naive_bayes import GaussianNB
# *** SVM classifier
from sklearn.svm import SVC
# --------  metrics:
from sklearn import metrics
from sklearn.metrics import accuracy_score, precision_score, recall_score, f1_score, confusion_matrix
from sklearn.metrics import make_scorer
```

<a id="dataset_desc"></a>
[go to basic data exploration](#data_exploration)
## The mushroom dataset
In this exercise you will work with the mushroom dataset.<br/>
**The mushroom dataset, describes two types of mushrooms**:<br/>
e - edible, p - poisonous.

**The following is a list of features in the dataset**:
1. cap-shape: bell=b,conical=c,convex=x,flat=f, knobbed=k,sunken=s
2. cap-surface: fibrous=f,grooves=g,scaly=y,smooth=s
3. cap-color: brown=n,buff=b,cinnamon=c,gray=g,green=r, pink=p,purple=u,red=e,white=w,yellow=y
4. bruises?: bruises=t,no=f
5. odor: almond=a,anise=l,creosote=c,fishy=y,foul=f, musty=m,none=n,pungent=p,spicy=s
6. gill-attachment: attached=a,descending=d,free=f,notched=n
7. gill-spacing: close=c,crowded=w,distant=d
8. gill-size: broad=b,narrow=n
9. gill-color: black=k,brown=n,buff=b,chocolate=h,gray=g, green=r,orange=o,pink=p,purple=u,red=e, white=w,yellow=y
10. stalk-shape: enlarging=e,tapering=t
11. stalk-root: bulbous=b,club=c,cup=u,equal=e, rhizomorphs=z,rooted=r,missing=?
12. stalk-surface-above-ring: fibrous=f,scaly=y,silky=k,smooth=s
13. stalk-surface-below-ring: fibrous=f,scaly=y,silky=k,smooth=s
14. stalk-color-above-ring: brown=n,buff=b,cinnamon=c,gray=g,orange=o, pink=p,red=e,white=w,yellow=y
15. stalk-color-below-ring: brown=n,buff=b,cinnamon=c,gray=g,orange=o, pink=p,red=e,white=w,yellow=y
16. veil-type: partial=p,universal=u
17. veil-color: brown=n,orange=o,white=w,yellow=y
18. ring-number: none=n,one=o,two=t
19. ring-type: cobwebby=c,evanescent=e,flaring=f,large=l, none=n,pendant=p,sheathing=s,zone=z
20. spore-print-color: black=k,brown=n,buff=b,chocolate=h,green=r, orange=o,purple=u,white=w,yellow=y
21. population: abundant=a,clustered=c,numerous=n, scattered=s,several=v,solitary=y
22. habitat: grasses=g,leaves=l,meadows=m,paths=p, urban=u,waste=w,woods=d

[go to basic data exploration](#data_exploration)

## 1. Load the dataset and prepare dataset for classification
In this section you will perform the following actions:
* 1.a. Load the mushrooms dataset
* 1.b. Basic data exploration
* 1.c. Prepare data for classification
* 1.d. Split dataset to train and test

### 1.a. Load the mushrooms dataset 
In this section you will load the mushrooms dataset from a csv file.

### 1.a. Instructions
<u>method name</u>: <b>load_dataset</b>
<pre>The following is expected:
--- Complete the 'load_dataset' function to load the mushroom dataset from the 'file_name' csv file
  into a pandas dataframe.
</pre>
<hr>
The return statement should look similar to the following statement:<br />
<b>return df_dataset</b>

```python
# 1.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ASSISTANCE TO ANSWER 
# ---- Add assistance code here IF NEEDED:
###
### YOUR CODE HERE
###

```

```python
# 1.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ANSWER 

def load_dataset(file_name):
   # טעינת קובץ ה-CSV של הפטריות
    df_dataset = pd.read_csv(file_name)
    
    return df_dataset
```

```python
# 1.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: SELF TESTING
# Use the following code to test your implementation:
file_name = '.' + os.sep + 'data' + os.sep + 'mushrooms.csv'
raw_dataset = load_dataset(file_name)
print(raw_dataset.describe())
# --- add additional code to check your code if needed:
print("\n--- חמש השורות הראשונות של הדאטה-סט ---")
print(raw_dataset.head())

```

Output:
```
class cap-shape cap-surface cap-color bruises  odor gill-attachment  \
count   8127      8127        8127      8126    8127  8127            8127   
unique     2         6           4        10       2     9               2   
top        e         x           y         n       f     n               f   
freq    4208      3657        3246      2285    4751  3528            7917   

       gill-spacing gill-size gill-color  ... stalk-surface-below-ring  \
count          8127      8127       8127  ...                     8127   
unique            2         2         12  ...                        4   
top               c         b          b  ...                        s   
freq           6815      5613       1730  ...                     4937   

       stalk-color-above-ring stalk-color-below-ring veil-type veil-color  \
count                    8127                   8127      8127       8127   
unique                      9                      9         1          4   
top           
```

```python
# 1.a. 
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation, if used)
# === CODE TYPE: SELF TESTING
# ---- Add your additional tests here if needed:
###
### YOUR CODE HERE
###

```

<!-- grade_id: test1a-1_load_dataset, points: 0.1, grade: true, solution: false -->
```python
# 1.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 1.a. - Test 1 (name: test1a-1_load_dataset, points: 0.1) - Sanity")
print ("\t--->Testing the implementation of 'load_dataset' ...")

file_name = '.' + os.sep + 'data' + os.sep + 'mushrooms.csv'

try:
    raw_dataset = load_dataset(file_name)
except Exception as e:
    print ('You probably have a syntax error, we got the following exception:')
    print ('\tError Message:', str(e))
    print ('Try fixing your implementation')
    raise 

print ("Good Job!\nYou've passed the 1st test for the 'load_dataset' function implementation :-)")
```

Output:
```
Part 1.a. - Test 1 (name: test1a-1_load_dataset, points: 0.1) - Sanity
	--->Testing the implementation of 'load_dataset' ...
Good Job!
You've passed the 1st test for the 'load_dataset' function implementation :-)
```

<!-- grade_id: test1a-2_load_dataset, points: 0.1, grade: true, solution: false -->
```python
# 1.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 1.a. - Test 2 (name: test1a-2_load_dataset, points: 0.1)")
print ("\t--->Testing the implementation of 'load_dataset' ...")

file_name = '.' + os.sep + 'data' + os.sep + 'mushrooms.csv'

try:
    raw_dataset = load_dataset(file_name)
except Exception as e:
    print ('You probably have a syntax error, we got the following exception:')
    print ('\tError Message:', str(e))
    print ('Try fixing your implementation')
    raise 

assert raw_dataset.shape == (8127, 23) , 'Wrong shape for dataset dataframe'

print ("Good Job!\nYou've passed the 2nd test for the 'load_dataset' function implementation :-)")
```

Output:
```
Part 1.a. - Test 2 (name: test1a-2_load_dataset, points: 0.1)
	--->Testing the implementation of 'load_dataset' ...
Good Job!
You've passed the 2nd test for the 'load_dataset' function implementation :-)
```

<a id="data_exploration"></a>
### 1.b. Basic data exploration
In this section you may perform simple processing on the dataset to understand it better.

* Perform simple dataset exploration. It is suggested to use describe, info and a simple histogram for features' values
* You could see also the [explanation of the dataset above](#dataset_desc)

```python
# 1.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: DATA EXPLORATION
# Perform a simple dataset exploration. 
# It is sugggested to use describe, info and a simple histogram for features' values
# Transfer dataset's values to numeric ones.
# Use the output of the exploration for next section.
# ---- Add assistance code here IF NEEDED:
# 1. הדפסת המידע המבני של ה-DataFrame (סוגי נתונים וערכים חסרים)
print("--- Dataset Info ---")
raw_dataset.info()

print("\n" + "="*50 + "\n")

# 2. בדיקה כמה ערכים ייחודיים יש בכל עמודה
print("--- Unique Values Per Feature ---")
for col in raw_dataset.columns:
    print(f"{col}: {raw_dataset[col].nunique()} unique values")

print("\n" + "="*50 + "\n")

# 3. בדיקת התפלגות עמודת המטרה (אכיל vs רעיל)
print("--- Target Class Distribution ---")
print(raw_dataset['class'].value_counts())
```

Output:
```
--- Dataset Info ---
<class 'pandas.core.frame.DataFrame'>
RangeIndex: 8127 entries, 0 to 8126
Data columns (total 23 columns):
 #   Column                    Non-Null Count  Dtype 
---  ------                    --------------  ----- 
 0   class                     8127 non-null   object
 1   cap-shape                 8127 non-null   object
 2   cap-surface               8127 non-null   object
 3   cap-color                 8126 non-null   object
 4   bruises                   8127 non-null   object
 5   odor                      8127 non-null   object
 6   gill-attachment           8127 non-null   object
 7   gill-spacing              8127 non-null   object
 8   gill-size                 8127 non-null   object
 9   gill-color                8127 non-null   object
 10  stalk-shape               8127 non-null   object
 11  stalk-root                8127 non-null   object
 12  stalk-surface-above-ring  8127 non-null   object
 13  stalk-surface-below-ring  8127 non-null   object
 14  sta
```

### 1.c. Prepare data for classification
In this section you will the <u>prepare data in the dataset for classification</u>.

* Use the results of your [data exploration above](#data_exploration)
* See also the [explanation of the dataset above](#dataset_desc)


### 1.c. Instructions
<u>method name</u>: <b>transfer_str_to_numeric_vals</b>
<pre>The following is expected:

- Remove any rows with one or more missing value. 
- For any duplicate rows, keep only the first one 
- Transfer dataset's values to numeric ones.
 Notes:
       Each unique string value should be transfered to a corresponding 
           unique numeric integer value.
       All the columns contain string values, as described above.
       It is suggested todo this in a generic way for all columns (you can do this part in two lines of code)
</pre>
<hr>
The return statement should look similar to the following statement:<br />
<b>return dataset</b>

```python
# 1.c.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ASSISTANCE TO ANSWER 
# ---- Add assistance code here IF NEEDED:
###
### YOUR CODE HERE
###

```

```python
# 1.c.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ANSWER 

def transfer_str_to_numeric_vals(dataset):
   # 1. הסרת שורות המכילות ערך חסר אחד או יותר
    processed_df = dataset.dropna()
    
    # 2. הסרת שורות כפולות ושמירה על המופע הראשון בלבד
    processed_df = processed_df.drop_duplicates(keep='first')
    
    # 3. המרה גנרית של כל עמודות הטקסט למספרים שלמים (Integers)
    # עבור כל עמודה, נמיר לטיפוס קטגורי וניקח את הייצוג המספרי שלו
    for col in processed_df.columns:
        processed_df[col] = processed_df[col].astype('category').cat.codes
        
    return processed_df
```

```python
# 1.c.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: SELF TESTING
# Use the following code to test your implementation:
file_name = '.' + os.sep + 'data' + os.sep + 'mushrooms.csv'
raw_dataset = load_dataset(file_name)
dataset = transfer_str_to_numeric_vals(raw_dataset)
cols = dataset.select_dtypes(include=[np.number])
arr_nums = [len(dataset[col].unique()) for col in cols]
print(arr_nums)
# --- add additional code to check your code if needed:
print("\n--- מידע על הדאטה-סט לאחר ניקוי והמרה ---")
print(f"מספר שורות סופי: {dataset.shape[0]}")
print(f"מספר עמודות סופי: {dataset.shape[1]}")
print("\nהצצה לחמש השורות הראשונות (כולן הפכו למספרים):")
print(dataset.head())
```

Output:
```
[2, 6, 4, 10, 2, 9, 2, 2, 2, 12, 2, 5, 4, 4, 9, 9, 1, 4, 3, 5, 9, 6, 7]

--- מידע על הדאטה-סט לאחר ניקוי והמרה ---
מספר שורות סופי: 8124
מספר עמודות סופי: 23

הצצה לחמש השורות הראשונות (כולן הפכו למספרים):
   class  cap-shape  cap-surface  cap-color  bruises  odor  gill-attachment  \
0      1          5            2          4        1     6                1   
1      0          5            2          9        1     0                1   
2      0          0            2          8        1     3                1   
3      1          5            3          8        1     6                1   
4      0          5            2          3        0     5                1   

   gill-spacing  gill-size  gill-color  ...  stalk-surface-below-ring  \
0             0          1           4  ...                         2   
1             0          0           4  ...                         2   
2             0          0           5  ...                         2   
3             0          1 
```

```python
# 1.c. 
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation, if used)
# === CODE TYPE: SELF TESTING
# ---- Add your additional tests here if needed:
###
### YOUR CODE HERE
###

```

<!-- grade_id: test1c-1_transfer_str_to_numeric_vals, points: 0.1, grade: true, solution: false -->
```python
# 1.c.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 1.c. - Test 1 (name: test1c-1_transfer_str_to_numeric_vals, points: 0.1) - Sanity")
print ("\t--->Testing the implementation of 'transfer_str_to_numeric_vals' ...")

file_name = '.' + os.sep + 'data' + os.sep + 'mushrooms.csv'

try:
    raw_dataset = load_dataset(file_name)
    dataset = transfer_str_to_numeric_vals(raw_dataset)
    cols = dataset.select_dtypes(include=[np.number])
    arr_nums = [len(dataset[col].unique()) for col in cols]
except Exception as e:
    print ('You probably have a syntax error, we got the following exception:')
    print ('\tError Message:', str(e))
    print ('Try fixing your implementation')
    raise 
    
print ("Good Job!\nYou've passed the 1st test for the 'transfer_str_to_numeric_vals' function implementation :-)")
```

Output:
```
Part 1.c. - Test 1 (name: test1c-1_transfer_str_to_numeric_vals, points: 0.1) - Sanity
	--->Testing the implementation of 'transfer_str_to_numeric_vals' ...
Good Job!
You've passed the 1st test for the 'transfer_str_to_numeric_vals' function implementation :-)
```

<!-- grade_id: test1c-2_transfer_str_to_numeric_vals, points: 0.1, grade: true, solution: false -->
```python
# 1.c.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 1.c. - Test 2 (name: test1c-2_transfer_str_to_numeric_vals, points: 0.1)")
print ("\t--->Testing the implementation of 'transfer_str_to_numeric_vals' ...")

file_name = '.' + os.sep + 'data' + os.sep + 'mushrooms.csv'

try:
    raw_dataset = load_dataset(file_name)
    dataset = transfer_str_to_numeric_vals(raw_dataset)
    cols = dataset.select_dtypes(include=[np.number])
    arr_nums = [len(dataset[col].unique()) for col in cols]
except Exception as e:
    print ('You probably have a syntax error, we got the following exception:')
    print ('\tError Message:', str(e))
    print ('Try fixing your implementation')
    raise 

assert dataset.shape == (8124, 23) , 'Wrong shape for dataset dataframe'
assert dataset.select_dtypes(include=[np.number]).columns.size == 23, 'Wrong number of numeric columns'

print ("Good Job!\nYou've passed the 2nd test for the 'transfer_str_to_numeric_vals' function implementation :-)")
```

Output:
```
Part 1.c. - Test 2 (name: test1c-2_transfer_str_to_numeric_vals, points: 0.1)
	--->Testing the implementation of 'transfer_str_to_numeric_vals' ...
Good Job!
You've passed the 2nd test for the 'transfer_str_to_numeric_vals' function implementation :-)
```

<!-- grade_id: test1c-3_transfer_str_to_numeric_vals, points: 0.2, grade: true, solution: false -->
```python
# 1.c.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 1.c. - Test 3 (name: test1c-3_transfer_str_to_numeric_vals, points: 0.2)")
print ("\t--->Testing the implementation of 'transfer_str_to_numeric_vals' ...")

file_name = '.' + os.sep + 'data' + os.sep + 'mushrooms.csv'

try:
    raw_dataset = load_dataset(file_name)
    dataset = transfer_str_to_numeric_vals(raw_dataset)
    cols = dataset.select_dtypes(include=[np.number])
    arr_nums = [len(dataset[col].unique()) for col in cols]
except Exception as e:
    print ('You probably have a syntax error, we got the following exception:')
    print ('\tError Message:', str(e))
    print ('Try fixing your implementation')
    raise 

assert np.round(np.mean(arr_nums),2)==5.17,'Wrong number of unique values'

print ("Good Job!\nYou've passed the 3rd test for the 'transfer_str_to_numeric_vals' function implementation :-)")
```

Output:
```
Part 1.c. - Test 3 (name: test1c-3_transfer_str_to_numeric_vals, points: 0.2)
	--->Testing the implementation of 'transfer_str_to_numeric_vals' ...
Good Job!
You've passed the 3rd test for the 'transfer_str_to_numeric_vals' function implementation :-)
```

### 1.d. Split dataset to train and test
In this section you will split the dataset into a train set and a test set.

### 1.d. Instructions
<u>method name</u>: <b>split_to_train_and_test</b>
<pre>The following is expected:
Step 1 - Split the given dataset into 'X' (feature vectors - dataframe)
      and 'y' (corresponding labels - series), determined by the given 'label_column' column.
Step 2 - Split X and y into 'X_train', 'X_test', and corresponding 'y_train' and 'y_test' series.
Notes: 
      The 'X_train' and 'X_test' dataframes should not include the 'label_column' column.
      Use sklearn's 'train_test_split' method, which was taught in class.
      Use the given 'rand_state' as the value for the 'random_state' parameter in 'train_test_split'.
</pre>
<hr>
The return statement should look similar to the following statement:<br />
<b>return X_train, X_test, y_train, y_test</b>

```python
# 1.d.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ASSISTANCE TO ANSWER 
# ---- Add assistance code here IF NEEDED:
###
### YOUR CODE HERE
###

```

```python
# 1.d.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ANSWER 

def split_to_train_and_test(dataset, label_column, test_ratio, rand_state):
   # שלב 1: הפרדת עמודת המטרה משאר המאפיינים
    X = dataset.drop(columns=[label_column])
    y = dataset[label_column]
    
    # שלב 2: חלוקת הנתונים תוך שימוש ביחס החלוקה וב-random state שקיבלנו
    X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=test_ratio, random_state=rand_state)
    
    return X_train, X_test, y_train, y_test
```

```python
# 1.d.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: SELF TESTING
# Use the following code to test your implementation:
file_name = '.' + os.sep + 'data' + os.sep + 'mushrooms.csv'
test_ratio, rand_state = 0.2, 42
category_col_name = 'class'
raw_dataset = load_dataset(file_name)
dataset = transfer_str_to_numeric_vals(raw_dataset) 
X_train, X_test, y_train, y_test = split_to_train_and_test(dataset, category_col_name, test_ratio, rand_state)
print(X_train.shape, X_test.shape, y_train.shape, y_test.shape)
# --- add additional code to check your code if needed:
print("\n--- וידוא אחוזי החלוקה ---")
total_rows = dataset.shape[0]
print(f"אחוז נתוני האימון: {X_train.shape[0] / total_rows:.1%}")
print(f"אחוז נתוני הטסט: {X_test.shape[0] / total_rows:.1%}")
```

Output:
```
(6499, 22) (1625, 22) (6499,) (1625,)

--- וידוא אחוזי החלוקה ---
אחוז נתוני האימון: 80.0%
אחוז נתוני הטסט: 20.0%
```

```python
# 1.d. 
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation, if used)
# === CODE TYPE: SELF TESTING
# ---- Add your additional tests here if needed:
###
### YOUR CODE HERE
###

```

<!-- grade_id: test1d-1_split_to_train_and_test, points: 0.1, grade: true, solution: false -->
```python
# 1.d.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 1.d. - Test 1 (name: test1d-1_split_to_train_and_test, points: 0.1) - Sanity")
print ("\t--->Testing the implementation of 'split_to_train_and_test' ...")

file_name = '.' + os.sep + 'data' + os.sep + 'mushrooms.csv'
test_ratio, rand_state = 0.2, 42
category_col_name = 'class'

try:
    raw_dataset = load_dataset(file_name)
    dataset = transfer_str_to_numeric_vals(raw_dataset) 
    X_train, X_test, y_train, y_test = split_to_train_and_test(dataset, category_col_name, test_ratio, rand_state)
except Exception as e:
    print ('You probably have a syntax error, we got the following exception:')
    print ('\tError Message:', str(e))
    print ('Try fixing your implementation')
    raise 

print ("Good Job!\nYou've passed the 1st test for the 'split_to_train_and_test' function implementation :-)")
```

Output:
```
Part 1.d. - Test 1 (name: test1d-1_split_to_train_and_test, points: 0.1) - Sanity
	--->Testing the implementation of 'split_to_train_and_test' ...
Good Job!
You've passed the 1st test for the 'split_to_train_and_test' function implementation :-)
```

<!-- grade_id: test1d-2_split_to_train_and_test, points: 0.1, grade: true, solution: false -->
```python
# 1.d.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 1.d. - Test 2 (name: test1d-2_split_to_train_and_test, points: 0.1)")
print ("\t--->Testing the implementation of 'split_to_train_and_test' ...")

file_name = '.' + os.sep + 'data' + os.sep + 'mushrooms.csv'
test_ratio, rand_state = 0.2, 42
category_col_name = 'class'

try:
    raw_dataset = load_dataset(file_name)
    dataset = transfer_str_to_numeric_vals(raw_dataset) 
    X_train, X_test, y_train, y_test = split_to_train_and_test(dataset, category_col_name, test_ratio, rand_state)
except Exception as e:
    print ('You probably have a syntax error, we got the following exception:')
    print ('\tError Message:', str(e))
    print ('Try fixing your implementation')
    raise 


assert X_train.shape == (6499, 22) and X_test.shape == (1625, 22), 'Wrong shape for feature-vector train or test dataframes'
assert y_train.shape[0] == 6499 and y_test.shape[0] == 1625, 'Wrong shape for category train or test series'

print ("Good Job!\nYou've passed the 2nd test for the 'split_to_train_and_test' function implementation :-)")
```

Output:
```
Part 1.d. - Test 2 (name: test1d-2_split_to_train_and_test, points: 0.1)
	--->Testing the implementation of 'split_to_train_and_test' ...
Good Job!
You've passed the 2nd test for the 'split_to_train_and_test' function implementation :-)
```

## 2. Auxiliary classification function
This section includes the following:
* 2.a. get classifier object 
* 2.b. get evaluation value on test 

### 2.a. get classifier object auxiliary function

The following 'get_classifier_obj' auxiliary function contain 2 input parameters:
- classifier_name
- params

And outputs the corresponding classifier

**The 'classifier_name' parameter**, contains a <u>string value</u><br />
   and indicates which classifier object to return, as following<br />
* if 'classifier_name' equals 'KNN' return 'KNeighborsClassifier()'
* if 'classifier_name' equals 'naive_bayes' return 'GaussianNB()'
* if 'classifier_name' equals 'svm' return 'SVC()'
* if 'classifier_name' equals 'decision_tree' return 'tree.DecisionTreeClassifier()'
* if 'classifier_name' equals 'random_forest ' return 'RandomForestClassifier()'

**<u>The 'param' parameter</u>** is a <u>dictionary object</u>.<br />
It could be 'None' or contain values as following:<br />
* If **'classifier_name' equals 'KNN'**:
    It will contain a value for the <u>'n_neighbors'</u> key in the 'param' parameter.
    * Set this ('n_neighbors') parameter in the returned 'KNeighborsClassifier' object accordingly.
* If **'classifier_name' equals 'decision_tree'**:
    It will contain a value for the <u>'max_depth' & 'min_samples_split'</u> keys in the 'param' parameter.
    * Set these ('max_depth', 'min_samples_split') parameters in the 'DecisionTreeClassifier' object accordingly.
* If **'classifier_name' equals 'random_forest'**:
    It will contain a value for the <u>'n_estimators'</u> key in the 'param' parameter.
    * Set this ('n_estimators') parameter in the 'RandomForestClassifier' object accordingly.
* If **'classifier_name' equals 'naive_bayes' or 'svm'**, assume the 'params' will be equal None.
    * In such a case return the relevant GaussianNB object or SVC object without setting any parameters.

### 2.a. Some examples

**Example 1**: The <u>'classifier_name' parameter equals 'KNN' and the 'param' parameter is not None</u>, <br />
return a classifier object as following:<br/ >
return KNeighborsClassifier(n_neighbors=params['n_neighbors'])

**Example 2**: The <u>'classifier_name' parameter equals 'KNN' and the 'param' parameter is None</u>,<br />
return a classifier object as following:<br />
return KNeighborsClassifier()

**Example 3**: The 'classifier_name' parameter equals 'svm',
return a classifier object as following:<br />
return SVC()

### 2.a. Instructions
<u>method name</u>: <b>get_classifier_obj</b>
<pre>The following is expected:
--- Complete the 'get_classifier_obj' auxiliary function to return
    a classifier object, as explained above.
</pre>
<hr>
The return statement should look similar to the following statement:<br />
<b>return clf</b>

```python
# 2.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ASSISTANCE TO ANSWER 
# ---- Add assistance code here IF NEEDED:
###
### YOUR CODE HERE
###

```

```python
# 2.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ANSWER 

def get_classifier_obj(classifier_name, params):
   # ניקוי רווחים מיותרים משם המסווג ליתר ביטחון
    name = classifier_name.strip()
    
    # 1. מסווג KNN
    if name == 'KNN':
        if params is not None and 'n_neighbors' in params:
            return KNeighborsClassifier(n_neighbors=params['n_neighbors'])
        return KNeighborsClassifier()
        
    # 2. מסווג עץ החלטה (Decision Tree)
    elif name == 'decision_tree':
        if params is not None:
            # שליפת הפרמטרים מתוך הדיקשנרי במידה וקיימים
            max_depth = params.get('max_depth')
            min_samples_split = params.get('min_samples_split')
            return DecisionTreeClassifier(max_depth=max_depth, min_samples_split=min_samples_split)
        return DecisionTreeClassifier()
        
    # 3. מסווג יער אקראי (Random Forest)
    elif name == 'random_forest':
        if params is not None and 'n_estimators' in params:
            return RandomForestClassifier(n_estimators=params['n_estimators'])
        return RandomForestClassifier()
        
    # 4. מסווג נייב באייס (Naive Bayes)
    elif name == 'naive_bayes':
        return GaussianNB()
        
    # 5. מסווג SVM
    elif name == 'svm':
        return SVC()
        
    # במקרה ששם המסווג לא מוכר (ליתר ביטחון)
    return None

```

```python
# 2.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: SELF TESTING
# Use the following code to test your implementation:
params_knn = {'n_neighbors':3}
params_random_forest = {'n_estimators':51}
params_decision_tree = {'max_depth':4, 'min_samples_split':4}
clf_naive_bayes = get_classifier_obj("naive_bayes",None)
clf_svm = get_classifier_obj("svm",None)
clf_knn = get_classifier_obj("KNN",None)
clf_random_forest = get_classifier_obj("random_forest",None)
clf_decision_tree = get_classifier_obj("decision_tree",None)
clf_knn_with_params = get_classifier_obj("KNN",params_knn)
clf_random_forest_with_params = get_classifier_obj("random_forest",params_random_forest)    
clf_decision_tree_with_params = get_classifier_obj("decision_tree",params_decision_tree)
# --- add additional code to check your code if needed:
print("--- בדיקת יצירת מסווגים ברירת מחדל ---")
print(f"Naive Bayes: {type(clf_naive_bayes)}")
print(f"SVM: {type(clf_svm)}")
print(f"KNN default: {clf_knn}")

print("\n--- בדיקת מסווגים עם פרמטרים מותאמים ---")
print(f"KNN (k=3): {clf_knn_with_params}")
print(f"Random Forest (estimators=51): {clf_random_forest_with_params}")
print(f"Decision Tree (depth=4): {clf_decision_tree_with_params}")
```

Output:
```
--- בדיקת יצירת מסווגים ברירת מחדל ---
Naive Bayes: <class 'sklearn.naive_bayes.GaussianNB'>
SVM: <class 'sklearn.svm._classes.SVC'>
KNN default: KNeighborsClassifier()

--- בדיקת מסווגים עם פרמטרים מותאמים ---
KNN (k=3): KNeighborsClassifier(n_neighbors=3)
Random Forest (estimators=51): RandomForestClassifier(n_estimators=51)
Decision Tree (depth=4): DecisionTreeClassifier(max_depth=4, min_samples_split=4)
```

```python
# 2.a. 
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation, if used)
# === CODE TYPE: SELF TESTING
# ---- Add your additional tests here if needed:
###
### YOUR CODE HERE
###

```

<!-- grade_id: test2a-1_get_classifier_obj, points: 0.1, grade: true, solution: false -->
```python
# 2.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 2.a. - Test 1 (name: test2a-1_get_classifier_obj, points: 0.1) - Sanity")
print ("\t--->Testing the implementation of 'get_classifier_obj' ...")


params_knn = {'n_neighbors':3}
params_random_forest = {'n_estimators':51}
params_decision_tree = {'max_depth':4, 'min_samples_split':4}
try:
    clf_naive_bayes = get_classifier_obj("naive_bayes",None)
    clf_svm = get_classifier_obj("svm",None)
    clf_knn = get_classifier_obj("KNN",None)
    clf_random_forest = get_classifier_obj("random_forest",None)
    clf_decision_tree = get_classifier_obj("decision_tree",None)
    clf_knn_with_params = get_classifier_obj("KNN",params_knn)
    clf_random_forest_with_params = get_classifier_obj("random_forest",params_random_forest)    
    clf_decision_tree_with_params = get_classifier_obj("decision_tree",params_decision_tree)
except Exception as e:
    print ('You probably have a syntax error, we got the following exception:')
    print ('\tError Message:', str(e))
    print ('Try fixing your implementation')
    raise 

print ("Good Job!\nYou've passed the 1st test for the 'get_classifier_obj' function :-)")
```

Output:
```
Part 2.a. - Test 1 (name: test2a-1_get_classifier_obj, points: 0.1) - Sanity
	--->Testing the implementation of 'get_classifier_obj' ...
Good Job!
You've passed the 1st test for the 'get_classifier_obj' function :-)
```

<!-- grade_id: test2a-2_get_classifier_obj, points: 0.1, grade: true, solution: false -->
```python
# 2.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 2.a. - Test 2 (name: test2a-2_get_classifier_obj, points: 0.1)")
print ("\t--->Testing the implementation of 'get_classifier_obj' ...")

try:
    clf = get_classifier_obj("naive_bayes",None)
except Exception as e:
    print ('You probably have a syntax error, we got the following exception:')
    print ('\tError Message:', str(e))
    print ('Try fixing your implementation')
    raise 

assert type(clf) == sklearn.naive_bayes.GaussianNB , 'Wrong type for classifier'

print ("Good Job!\nYou've passed the 2nd test for the 'get_classifier_obj' (for 'Naive Nayes') function :-)")
```

Output:
```
Part 2.a. - Test 2 (name: test2a-2_get_classifier_obj, points: 0.1)
	--->Testing the implementation of 'get_classifier_obj' ...
Good Job!
You've passed the 2nd test for the 'get_classifier_obj' (for 'Naive Nayes') function :-)
```

<!-- grade_id: test2a-3_get_classifier_obj, points: 0.1, grade: true, solution: false -->
```python
# 2.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 2.a. - Test 3 (name: test2a-3_get_classifier_obj, points: 0.1)")
print ("\t--->Testing the implementation of 'get_classifier_obj' ...")

try:
    clf = get_classifier_obj("svm",None)
except Exception as e:
    print ('You probably have a syntax error, we got the following exception:')
    print ('\tError Message:', str(e))
    print ('Try fixing your implementation')
    raise 

assert type(clf) == sklearn.svm._classes.SVC , 'Wrong type for classifier'

print ("Good Job!\nYou've passed the 3rd test for the 'get_classifier_obj' (for 'SVM') function :-)")
```

Output:
```
Part 2.a. - Test 3 (name: test2a-3_get_classifier_obj, points: 0.1)
	--->Testing the implementation of 'get_classifier_obj' ...
Good Job!
You've passed the 3rd test for the 'get_classifier_obj' (for 'SVM') function :-)
```

<!-- grade_id: test2a-4i_get_classifier_obj, points: 0.1, grade: true, solution: false -->
```python
# 2.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 2.a. - Test 4i (name: test2a-4i_get_classifier_obj, points: 0.1)")
print ("\t--->Testing the implementation of 'get_classifier_obj' ...")

try:
    clf = get_classifier_obj("KNN",None)
except Exception as e:
    print ('You probably have a syntax error, we got the following exception:')
    print ('\tError Message:', str(e))
    print ('Try fixing your implementation')
    raise 

assert type(clf) == sklearn.neighbors._classification.KNeighborsClassifier , 'Wrong type for classifier'
assert clf.n_neighbors == 5 , 'Wrong value for n_neighbors in KNN, expected default value'

print ("Good Job!\nYou've passed the 4th test (without paramaters) for the 'get_classifier_obj' (for 'KNN') function :-)")
```

Output:
```
Part 2.a. - Test 4i (name: test2a-4i_get_classifier_obj, points: 0.1)
	--->Testing the implementation of 'get_classifier_obj' ...
Good Job!
You've passed the 4th test (without paramaters) for the 'get_classifier_obj' (for 'KNN') function :-)
```

<!-- grade_id: test2a-4ii_get_classifier_obj, points: 0.1, grade: true, solution: false -->
```python
# 2.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 2.a. - Test 4ii (name: test2a-4ii_get_classifier_obj, points: 0.1)")
print ("\t--->Testing the implementation of 'get_classifier_obj' ...")

params = {'n_neighbors':3}
try:
    clf = get_classifier_obj("KNN",params)
except Exception as e:
    print ('You probably have a syntax error, we got the following exception:')
    print ('\tError Message:', str(e))
    print ('Try fixing your implementation')
    raise 

assert type(clf) == sklearn.neighbors._classification.KNeighborsClassifier , 'Wrong type for classifier'
assert clf.n_neighbors == 3 , 'Wrong value for n_neighbors in KNN, expected to be equal to input value'

print ("Good Job!\nYou've passed the 4th test (with paramaters) for the 'get_classifier_obj' (for 'KNN') function :-)")
```

Output:
```
Part 2.a. - Test 4ii (name: test2a-4ii_get_classifier_obj, points: 0.1)
	--->Testing the implementation of 'get_classifier_obj' ...
Good Job!
You've passed the 4th test (with paramaters) for the 'get_classifier_obj' (for 'KNN') function :-)
```

<!-- grade_id: test2a-5i_get_classifier_obj, points: 0.1, grade: true, solution: false -->
```python
# 2.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 2.a. - Test 5i (name: test2a-5i_get_classifier_obj, points: 0.1)")
print ("\t--->Testing the implementation of 'get_classifier_obj' ...")

try:
    clf = get_classifier_obj("random_forest",None)
except Exception as e:
    print ('You probably have a syntax error, we got the following exception:')
    print ('\tError Message:', str(e))
    print ('Try fixing your implementation')
    raise 

assert type(clf) == sklearn.ensemble._forest.RandomForestClassifier , 'Wrong type for classifier'
assert clf.n_estimators == 100 , 'Wrong value for n_estimators in random forest, expected default value'

print ("Good Job!\nYou've passed the 5th test (without paramaters) for the 'get_classifier_obj' (for' Random Forest') function :-)")
```

Output:
```
Part 2.a. - Test 5i (name: test2a-5i_get_classifier_obj, points: 0.1)
	--->Testing the implementation of 'get_classifier_obj' ...
Good Job!
You've passed the 5th test (without paramaters) for the 'get_classifier_obj' (for' Random Forest') function :-)
```

<!-- grade_id: test2a-5ii_get_classifier_obj, points: 0.1, grade: true, solution: false -->
```python
# 2.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 2.a. - Test 5ii (name: test2a-5ii_get_classifier_obj, points: 0.1)")
print ("\t--->Testing the implementation of 'get_classifier_obj' ...")

params = {'n_estimators':51}

try:
    clf = get_classifier_obj("random_forest",params)
except Exception as e:
    print ('You probably have a syntax error, we got the following exception:')
    print ('\tError Message:', str(e))
    print ('Try fixing your implementation')
    raise 

assert type(clf) == sklearn.ensemble._forest.RandomForestClassifier , 'Wrong type for classifier'
assert clf.n_estimators == 51 , 'Wrong value for n_estimators in random forest, expected to be equal to input value'

print ("Good Job!\nYou've passed the 5th test (with paramaters) for the 'get_classifier_obj' (for' Random Forest') function :-)")
```

Output:
```
Part 2.a. - Test 5ii (name: test2a-5ii_get_classifier_obj, points: 0.1)
	--->Testing the implementation of 'get_classifier_obj' ...
Good Job!
You've passed the 5th test (with paramaters) for the 'get_classifier_obj' (for' Random Forest') function :-)
```

<!-- grade_id: test2a-6i_get_classifier_obj, points: 0.1, grade: true, solution: false -->
```python
# 2.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 2.a. - Test 6i (name: test2a-6i_get_classifier_obj, points: 0.1)")
print ("\t--->Testing the implementation of 'get_classifier_obj' ...")

try:
    clf = get_classifier_obj("decision_tree",None)
except Exception as e:
    print ('You probably have a syntax error, we got the following exception:')
    print ('\tError Message:', str(e))
    print ('Try fixing your implementation')
    raise 

assert type(clf) == sklearn.tree._classes.DecisionTreeClassifier , 'Wrong type for classifier'
assert clf.max_depth is None and clf.min_samples_split==2 , 'Wrong values for max_depth or min_samples_split in Decision Tree, expected default value'

print ("Good Job!\nYou've passed the 6th test (without paramaters) for the 'get_classifier_obj' (for 'Decision Tree') function :-)")
```

Output:
```
Part 2.a. - Test 6i (name: test2a-6i_get_classifier_obj, points: 0.1)
	--->Testing the implementation of 'get_classifier_obj' ...
Good Job!
You've passed the 6th test (without paramaters) for the 'get_classifier_obj' (for 'Decision Tree') function :-)
```

<!-- grade_id: test2a-6ii_get_classifier_obj, points: 0.1, grade: true, solution: false -->
```python
# 2.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 2.a. - Test 6ii (name: test2a-6ii_get_classifier_obj, points: 0.1)")
print ("\t--->Testing the implementation of 'get_classifier_obj' ...")

params = {'max_depth':4, 'min_samples_split':4}
try:
    clf = get_classifier_obj("decision_tree",params)
except Exception as e:
    print ('You probably have a syntax error, we got the following exception:')
    print ('\tError Message:', str(e))
    print ('Try fixing your implementation')
    raise 

assert type(clf) == sklearn.tree._classes.DecisionTreeClassifier , 'Wrong type for classifier'
assert clf.max_depth==4 and clf.min_samples_split==4 , 'Wrong values for max_depth or min_samples_split in Decision Tree, expected to be equal to input values'

print ("Good Job!\nYou've passed the 6th test (with paramaters) for the 'get_classifier_obj' (for 'Decision Tree') function :-)")
```

Output:
```
Part 2.a. - Test 6ii (name: test2a-6ii_get_classifier_obj, points: 0.1)
	--->Testing the implementation of 'get_classifier_obj' ...
Good Job!
You've passed the 6th test (with paramaters) for the 'get_classifier_obj' (for 'Decision Tree') function :-)
```

### 2.b. get evaluation value on test auxiliary function

The following **'calc_evaluation_val'** auxiliary function contain <u>3 input parameters</u>:
- eval_metric - the evaluation metric expected to return
- y_test      - the actual test target categories
- y_predicted - the predicted test categories

**The 'eval_metric' parameter**, contains a <u>string value</u><br />
   and indicates which evaluation metric, which object to return, as following:<br />
* if 'eval_metric' equals <u>'accuracy'</u> return the <u>accuracy_score</u> evaluation float value
* if 'eval_metric' equals <u>'precision'</u> return the <u>precision_score</u> evaluation float value
* if 'eval_metric' equals <u>'recall'</u> return the <u>recall_score</u> evaluation float value
* if 'eval_metric' equals <u>'f1'</u> return the <u>f1_score</u> evaluation float value
* if 'eval_metric' equals <u>'confusion_matrix'</u> return the <u>confusion_matrix</u> evaluation np.ndarray value 

### 2.b. Instructions
<u>method name</u>: <b>calc_evaluation_val</b>
<pre>The following is expected:
--- Complete the 'calc_evaluation_val' auxiliary function to return
    the evaluation value of the trained model, as explained above.
</pre>
<hr>
The return statement should look similar to the following statement:<br />
<b>return evaluation_val</b>

```python
# 2.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ASSISTANCE TO ANSWER 
# ---- Add assistance code here IF NEEDED:
###
### YOUR CODE HERE
###

```

```python
# 2.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ANSWER 

def calc_evaluation_val(eval_metric, y_test, y_predicted):
   # 1. חישוב Accuracy
    if eval_metric == 'accuracy':
        return accuracy_score(y_test, y_predicted)
        
    # 2. חישוב Precision
    elif eval_metric == 'precision':
        return precision_score(y_test, y_predicted)
        
    # 3. חישוב Recall
    elif eval_metric == 'recall':
        return recall_score(y_test, y_predicted)
        
    # 4. חישוב F1-Score
    elif eval_metric == 'f1':
        return f1_score(y_test, y_predicted)
        
    # 5. חישוב מטריצת בלבול (Confusion Matrix)
    elif eval_metric == 'confusion_matrix':
        return confusion_matrix(y_test, y_predicted)
        
    # ליתר ביטחון, אם הועבר מדד לא מוכר
    return None
```

```python
# 2.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: SELF TESTING
# Use the following code to test your implementation:
y_test = pd.Series([1,0,1,0,1,0,1,0,1,0])
y_predicted = pd.Series([1,1,1,1,1,1,1,0,0,0])
accuracy_val = calc_evaluation_val("accuracy", y_test, y_predicted)
precision_val = calc_evaluation_val("precision", y_test, y_predicted)
recall_val = calc_evaluation_val("recall", y_test, y_predicted)
f1_val = calc_evaluation_val("f1", y_test, y_predicted)
confusion_matrix_val = calc_evaluation_val("confusion_matrix", y_test, y_predicted)
# --- add additional code to check your code if needed:
print("--- תוצאות מדדי ההערכה (בדיקה עצמית) ---")
print(f"Accuracy:  {accuracy_val}")
print(f"Precision: {precision_val:.3f}")
print(f"Recall:    {recall_val:.3f}")
print(f"F1 Score:  {f1_val:.3f}")
print("\nConfusion Matrix:")
print(confusion_matrix_val)
```

Output:
```
--- תוצאות מדדי ההערכה (בדיקה עצמית) ---
Accuracy:  0.6
Precision: 0.571
Recall:    0.800
F1 Score:  0.667

Confusion Matrix:
[[2 3]
 [1 4]]
```

```python
# 2.b. 
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation, if used)
# === CODE TYPE: SELF TESTING
# ---- Add your additional tests here if needed:
###
### YOUR CODE HERE
###

```

<!-- grade_id: test2b-1_calc_evaluation_val, points: 0.1, grade: true, solution: false -->
```python
# 2.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 2.b. - Test 1 (name: test2b-1_calc_evaluation_val, points: 0.1) - Sanity")
print ("\t--->Testing the implementation of 'calc_evaluation_val' ...")

try:
    y_test = pd.Series([1,0,1,0,1,0,1,0,1,0])
    y_predicted = pd.Series([1,1,1,1,1,1,1,0,0,0])
    accuracy_val = calc_evaluation_val("accuracy", y_test, y_predicted)
    precision_val = calc_evaluation_val("precision", y_test, y_predicted)
    recall_val = calc_evaluation_val("recall", y_test, y_predicted)
    f1_val = calc_evaluation_val("f1", y_test, y_predicted)
    confusion_matrix_val = calc_evaluation_val("confusion_matrix", y_test, y_predicted)
except Exception as e:
    print ('You probably have a syntax error, we got the following exception:')
    print ('\tError Message:', str(e))
    print ('Try fixing your implementation')
    raise 

print ("Good Job!\nYou've passed the 1st test for the 'calc_evaluation_val' function implementation :-)")
```

Output:
```
Part 2.b. - Test 1 (name: test2b-1_calc_evaluation_val, points: 0.1) - Sanity
	--->Testing the implementation of 'calc_evaluation_val' ...
Good Job!
You've passed the 1st test for the 'calc_evaluation_val' function implementation :-)
```

<!-- grade_id: test2b-2_calc_evaluation_val, points: 0.1, grade: true, solution: false -->
```python
# 2.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 2.b. - Test 2 (name: test2b-2_calc_evaluation_val, points: 0.1)")
print ("\t--->Testing the implementation of 'calc_evaluation_val' ...")

try:
    y_test = pd.Series([1,0,1,0,1,0,1,0,1,0])
    y_predicted = pd.Series([1,1,1,1,1,1,1,0,0,0])
    evaluation_val = calc_evaluation_val("accuracy", y_test, y_predicted)
except Exception as e:
    print ('You probably have a syntax error, we got the following exception:')
    print ('\tError Message:', str(e))
    print ('Try fixing your implementation')
    raise 
    
assert 0.6 == np.round(evaluation_val,2),"Wrong 'accuracy' value"

print("accuracy:",evaluation_val)

print ("Good Job!\nYou've passed the 2nd test for the 'calc_evaluation_val' function implementation (for 'accuracy') :-)")
```

Output:
```
Part 2.b. - Test 2 (name: test2b-2_calc_evaluation_val, points: 0.1)
	--->Testing the implementation of 'calc_evaluation_val' ...
accuracy: 0.6
Good Job!
You've passed the 2nd test for the 'calc_evaluation_val' function implementation (for 'accuracy') :-)
```

<!-- grade_id: test2b-3_calc_evaluation_val, points: 0.1, grade: true, solution: false -->
```python
# 2.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 2.b. - Test 3 (name: test2b-3_calc_evaluation_val, points: 0.1)")
print ("\t--->Testing the implementation of 'calc_evaluation_val' ...")

try:
    y_test = pd.Series([1,0,1,0,1,0,1,0,1,0])
    y_predicted = pd.Series([1,1,1,1,1,1,1,0,0,0])
    evaluation_val = calc_evaluation_val("precision", y_test, y_predicted)
except Exception as e:
    print ('You probably have a syntax error, we got the following exception:')
    print ('\tError Message:', str(e))
    print ('Try fixing your implementation')
    raise 

assert 0.57 == np.round(evaluation_val,2),"Wrong 'precision' value"

print("precision:",evaluation_val)

print ("Good Job!\nYou've passed the 3rd test for the 'calc_evaluation_val' function implementation (for 'precision') :-)")
```

Output:
```
Part 2.b. - Test 3 (name: test2b-3_calc_evaluation_val, points: 0.1)
	--->Testing the implementation of 'calc_evaluation_val' ...
precision: 0.5714285714285714
Good Job!
You've passed the 3rd test for the 'calc_evaluation_val' function implementation (for 'precision') :-)
```

<!-- grade_id: test2b-4_calc_evaluation_val, points: 0.1, grade: true, solution: false -->
```python
# 2.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 2.b. - Test 4 (name: test2b-4_calc_evaluation_val, points: 0.1)")
print ("\t--->Testing the implementation of 'calc_evaluation_val' ...")

try:
    y_test = pd.Series([1,0,1,0,1,0,1,0,1,0])
    y_predicted = pd.Series([1,1,1,1,1,1,1,0,0,0])
    evaluation_val = calc_evaluation_val("recall", y_test, y_predicted)
except Exception as e:
    print ('You probably have a syntax error, we got the following exception:')
    print ('\tError Message:', str(e))
    print ('Try fixing your implementation')
    raise 
    
assert 0.8 == np.round(evaluation_val,2),"Wrong 'recall' value"

print("recall:",evaluation_val)

print ("Good Job!\nYou've passed the 4th test for the 'calc_evaluation_val' function implementation (for 'recall') :-)")
```

Output:
```
Part 2.b. - Test 4 (name: test2b-4_calc_evaluation_val, points: 0.1)
	--->Testing the implementation of 'calc_evaluation_val' ...
recall: 0.8
Good Job!
You've passed the 4th test for the 'calc_evaluation_val' function implementation (for 'recall') :-)
```

<!-- grade_id: test2b-5_calc_evaluation_val, points: 0.1, grade: true, solution: false -->
```python
# 2.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 2.b. - Test 5 (name: test2b-5_calc_evaluation_val, points: 0.1)")
print ("\t--->Testing the implementation of 'calc_evaluation_val' ...")

try:
    y_test = pd.Series([1,0,1,0,1,0,1,0,1,0])
    y_predicted = pd.Series([1,1,1,1,1,1,1,0,0,0])
    evaluation_val = calc_evaluation_val("f1", y_test, y_predicted)
except Exception as e:
    print ('You probably have a syntax error, we got the following exception:')
    print ('\tError Message:', str(e))
    print ('Try fixing your implementation')
    raise 

assert 0.67 == np.round(evaluation_val,2),"Wrong 'f1' value"

print("f1:",evaluation_val)

print ("Good Job!\nYou've passed the 5th test for the 'calc_evaluation_val' function implementation (for 'f1') :-)")
```

Output:
```
Part 2.b. - Test 5 (name: test2b-5_calc_evaluation_val, points: 0.1)
	--->Testing the implementation of 'calc_evaluation_val' ...
f1: 0.6666666666666666
Good Job!
You've passed the 5th test for the 'calc_evaluation_val' function implementation (for 'f1') :-)
```

<!-- grade_id: test2b-6_calc_evaluation_val, points: 0.1, grade: true, solution: false -->
```python
# 2.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 2.b. - Test 6 (name: test2b-6_calc_evaluation_val, points: 0.1)")
print ("\t--->Testing the implementation of 'calc_evaluation_val' ...")

try:
    y_test = pd.Series([1,0,1,0,1,0,1,0,1,0])
    y_predicted = pd.Series([1,1,1,1,1,1,1,0,0,0])
    evaluation_val = calc_evaluation_val("confusion_matrix", y_test, y_predicted)
except Exception as e:
    print ('You probably have a syntax error, we got the following exception:')
    print ('\tError Message:', str(e))
    print ('Try fixing your implementation')
    raise 

assert np.ndarray==type(evaluation_val),'Wrong evaluation value for confusion matrix'

assert 1 == evaluation_val[1,0],"Wrong 'False Negative' value"

print("confusion_matrix:\n---------------------")
df_confusion_matrix = pd.DataFrame(evaluation_val,index=['Actual Negative','Actual Positive'],columns=['Predicted Negative','Predicted Positive'])
print(df_confusion_matrix)

print ("Good Job!\nYou've passed the 6th test for the 'calc_evaluation_val' function implementation (for 'confusion matrix') :-)")
```

Output:
```
Part 2.b. - Test 6 (name: test2b-6_calc_evaluation_val, points: 0.1)
	--->Testing the implementation of 'calc_evaluation_val' ...
confusion_matrix:
---------------------
                 Predicted Negative  Predicted Positive
Actual Negative                   2                   3
Actual Positive                   1                   4
Good Job!
You've passed the 6th test for the 'calc_evaluation_val' function implementation (for 'confusion matrix') :-)
```

## 3. Grid Search Cross-validation - hyper-parameter tuning

### 3.a. KNN - find best K for KNN

### 3.a. Instructions
<u>method name</u>: <b>find_best_k_for_KNN</b>
<pre>The following is expected:
--- Complete the 'find_best_k_for_KNN' function to find the K-hyperparameter (i.e. the 
     'n_neighbors' parameter) for the KNN algorithm, running on the mushroom 
     dataset's train-set, using grid search cross validation 'GridSearchCV' function.
Notes:
      You need to use the average f1 score, in order to choose the best K 
         for this dataset.
      Use the 3,7,9,11 possible values as possibilities to the 'n_neighbors' parameter.
      You need to return the best K value as well as the best average f1 score.
      Use the 'make_scorer', as the function for 'scoring' parameter
</pre>      
<hr>
The return statement should look similar to the following statement:<br />
<b>return best_K, best_f1_val</b>

```python
# 3.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ASSISTANCE TO ANSWER 
# ---- Add assistance code here IF NEEDED:
###
### YOUR CODE HERE
###

```

```python
# 3.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ANSWER 

def find_best_k_for_KNN(X_train, y_train):
   # 1. הגדרת רשת הפרמטרים שברצוננו לבדוק
    param_grid = {'n_neighbors': [3, 7, 9, 11]}
    
    # 2. יצירת אובייקט ה-Scorer עבור מדד f1
    f1_scorer = make_scorer(f1_score)
    
    # 3. יצירת מודל KNN בסיסי
    knn = KNeighborsClassifier()
    
    # 4. הגדרת ה-GridSearchCV (כברירת מחדל משתמש ב-5-Fold Cross Validation)
    grid_search = GridSearchCV(estimator=knn, param_grid=param_grid, scoring=f1_scorer)
    
    # 5. הרצת החיפוש והאימון על נתוני האימון
    grid_search.fit(X_train, y_train)
    
    # 6. שליפת ה-K המנצח והציון הממוצע הכי גבוה שלו
    best_K = grid_search.best_params_['n_neighbors']
    best_f1_val = grid_search.best_score_
    
    return best_K, best_f1_val
```

```python
# 3.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: SELF TESTING
# Use the following code to test your implementation:
file_name = '.' + os.sep + 'data' + os.sep + 'mushrooms.csv'
test_ratio, rand_state = 0.2, 42
category_col_name = 'class'
raw_dataset = load_dataset(file_name)
sub_set_of_raw_dataset_2k_4k = raw_dataset.iloc[2000:4000,:]
dataset = transfer_str_to_numeric_vals(sub_set_of_raw_dataset_2k_4k) 
X_train, X_test, y_train, y_test = split_to_train_and_test(dataset, category_col_name, test_ratio, rand_state)
best_K, best_f1_KNN_params = find_best_k_for_KNN(X_train, y_train)
# --- add additional code to check your code if needed:
print("--- תוצאות GridSearchCV עבור KNN ---")
print(f"ה-K הטוב ביותר שנבחר: {best_K}")
print(f"ציון ה-F1 הממוצע הטוב ביותר: {best_f1_KNN_params:.4f}")
```

Output:
```
--- תוצאות GridSearchCV עבור KNN ---
ה-K הטוב ביותר שנבחר: 7
ציון ה-F1 הממוצע הטוב ביותר: 0.9932
```

```python
# 3.a. 
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation, if used)
# === CODE TYPE: SELF TESTING
# ---- Add your additional tests here if needed:
###
### YOUR CODE HERE
###

```

<!-- grade_id: test3a-1_find_best_k_for_KNN, points: 0.1, grade: true, solution: false -->
```python
# 3.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 3.a. - Test 1 (name: test3a-1_find_best_k_for_KNN, points: 0.1) - Sanity")
print ("\t--->Testing the implementation of 'find_best_k_for_KNN' ...")

file_name = '.' + os.sep + 'data' + os.sep + 'mushrooms.csv'
test_ratio, rand_state = 0.2, 42
category_col_name = 'class'

try:
    raw_dataset = load_dataset(file_name)
    sub_set_of_raw_dataset_2k_4k = raw_dataset.iloc[2000:4000,:]
    dataset = transfer_str_to_numeric_vals(sub_set_of_raw_dataset_2k_4k) 
    X_train, X_test, y_train, y_test = split_to_train_and_test(dataset, category_col_name, test_ratio, rand_state)
    best_K, best_f1_KNN_params = find_best_k_for_KNN(X_train, y_train)
except Exception as e:
    print ('You probably have a syntax error, we got the following exception:')
    print ('\tError Message:', str(e))
    print ('Try fixing your implementation')
    raise 

print ("Good Job!\nYou've passed the 1st test for the 'find_best_k_for_KNN' function implementation :-)")
```

Output:
```
Part 3.a. - Test 1 (name: test3a-1_find_best_k_for_KNN, points: 0.1) - Sanity
	--->Testing the implementation of 'find_best_k_for_KNN' ...
Good Job!
You've passed the 1st test for the 'find_best_k_for_KNN' function implementation :-)
```

<!-- grade_id: test3a-2_find_best_k_for_KNN, points: 0.4, grade: true, solution: false -->
```python
# 3.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 3.a. - Test 2 (name: test3a-2_find_best_k_for_KNN, points: 0.4)")
print ("\t--->Testing the implementation of 'find_best_k_for_KNN' ...")

file_name = '.' + os.sep + 'data' + os.sep + 'mushrooms.csv'
test_ratio, rand_state = 0.2, 42
category_col_name = 'class'

try:
    raw_dataset = load_dataset(file_name)
    sub_set_of_raw_dataset_2k_4k = raw_dataset.iloc[2000:4000,:]
    dataset = transfer_str_to_numeric_vals(sub_set_of_raw_dataset_2k_4k) 
    X_train, X_test, y_train, y_test = split_to_train_and_test(dataset, category_col_name, test_ratio, rand_state)
except Exception as e:
    print ('You probably have a syntax error, we got the following exception:')
    print ('\tError Message:', str(e))
    print ('Try fixing your implementation')
    raise 

found_assertion_combination = False    
for i in range(50):
    try:
        best_K, best_f1_KNN_params = find_best_k_for_KNN(X_train, y_train)
        ###
        ### AUTOGRADER TEST - DO NOT REMOVE
        ###
        assert best_K==7,'Wrong best K hyperparameter value for KNN'
        found_assertion_combination = True
        print ('... successful')
        break
    except Exception as e:
        print ('Try:',i+1,'failed')
        print ('\tError Message:', str(e))
        print ('Will try again ...')

assert found_assertion_combination,'Wrong best K hyperparameter value for KNN'

print ("Good Job!\nYou've passed the 2nd test for the 'find_best_k_for_KNN' function implementation :-)")
```

Output:
```
Part 3.a. - Test 2 (name: test3a-2_find_best_k_for_KNN, points: 0.4)
	--->Testing the implementation of 'find_best_k_for_KNN' ...
... successful
Good Job!
You've passed the 2nd test for the 'find_best_k_for_KNN' function implementation :-)
```

<!-- grade_id: test3a-3_find_best_k_for_KNN, points: 0.4, grade: true, solution: false -->
```python
# 3.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 3.a. - Test 3 (name: test3a-3_find_best_k_for_KNN, points: 0.4)")
print ("\t--->Testing the implementation of 'find_best_k_for_KNN' ...")

file_name = '.' + os.sep + 'data' + os.sep + 'mushrooms.csv'
test_ratio, rand_state = 0.2, 42
category_col_name = 'class'

try:
    raw_dataset = load_dataset(file_name)
    sub_set_of_raw_dataset_2k_4k = raw_dataset.iloc[2000:4000,:]
    dataset = transfer_str_to_numeric_vals(sub_set_of_raw_dataset_2k_4k) 
    X_train, X_test, y_train, y_test = split_to_train_and_test(dataset, category_col_name, test_ratio, rand_state)
except Exception as e:
    print ('You probably have a syntax error, we got the following exception:')
    print ('\tError Message:', str(e))
    print ('Try fixing your implementation')
    raise 

found_assertion_combination = False
for i in range(50):
    try:
        best_K, best_f1_KNN_params = find_best_k_for_KNN(X_train, y_train)
        ###
        ### AUTOGRADER TEST - DO NOT REMOVE
        ###
        assert 0.993==np.round(best_f1_KNN_params,3),'Wrong F1 value for best K (number of neigbors) for KNN'        
        found_assertion_combination = True
        print ('... successful')
        break
    except Exception as e:
        print ('Try:',i+1,'failed')
        print ('\tError Message:', str(e))
        print ('Will try again ...')

assert found_assertion_combination, 'Wrong F1 value for best K (number of neigbors) for KNN'

print ("Good Job!\nYou've passed the 3rd test for the 'find_best_k_for_KNN' function implementation :-)")
```

Output:
```
Part 3.a. - Test 3 (name: test3a-3_find_best_k_for_KNN, points: 0.4)
	--->Testing the implementation of 'find_best_k_for_KNN' ...
... successful
Good Job!
You've passed the 3rd test for the 'find_best_k_for_KNN' function implementation :-)
```

### 3.b. Decision tree - find best parameter pair

### 3.b. Instructions
<u>method name</u>: <b>find_best_decision_tree_params</b>
<pre>The following is expected:
--- Complete the 'find_best_decision_tree_params' function to find best
          'max_depth' and 'min_samples_split' hyperparameter pair,
          for the decision-trees algorithm, running on the mushroom dataset.
     Again, use the grid search cross validation 'GridSearchCV' function.
Notes:
       You need to use the average f1 score, in order to choose the best  
           'max_depth' and 'min_samples_split' hyperparameter pair for this dataset.
       Use 2,4,6 as possible values as possibilities for the 'max_depth' parameter
           and 5,10,20 for 'min_samples_split'.
       You need to return the value of best 'max_depth', 'min_samples_split' pair
           and best average f1 score.
       Use the 'make_scorer', as the function for 'scoring' parameter
</pre>
<hr>
The return statement should look similar to the following statement:<br />
<b>return best_max_depth, best_min_samples_split, best_f1_val</b>

```python
# 3.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ASSISTANCE TO ANSWER 
# ---- Add assistance code here IF NEEDED:
###
### YOUR CODE HERE
###

```

```python
# 3.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ANSWER 

def find_best_decision_tree_params(X_train, y_train):
   # 1. הגדרת רשת הפרמטרים עבור עץ ההחלטה
    param_grid = {
        'max_depth': [2, 4, 6],
        'min_samples_split': [5, 10, 20]
    }
    
    # 2. יצירת אובייקט ה-Scorer עבור מדד f1
    f1_scorer = make_scorer(f1_score)
    
    # 3. יצירת מודל עץ החלטה בסיסי
    dt_clf = DecisionTreeClassifier()
    
    # 4. הגדרת ה-GridSearchCV (משתמש ב-5-Fold Cross Validation כברירת מחדל)
    grid_search = GridSearchCV(estimator=dt_clf, param_grid=param_grid, scoring=f1_scorer)
    
    # 5. הרצת החיפוש והאימון על נתוני האימון
    grid_search.fit(X_train, y_train)
    
    # 6. שליפת הערכים המנצחים והציון הממוצע הכי גבוה
    best_max_depth = grid_search.best_params_['max_depth']
    best_min_samples_split = grid_search.best_params_['min_samples_split']
    best_f1_val = grid_search.best_score_
    
    return best_max_depth, best_min_samples_split, best_f1_val
```

```python
# 3.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: SELF TESTING
# Use the following code to test your implementation:
file_name = '.' + os.sep + 'data' + os.sep + 'mushrooms.csv'
test_ratio, rand_state = 0.2, 42
category_col_name = 'class'
raw_dataset = load_dataset(file_name)
sub_set_of_raw_dataset_2k_4k = raw_dataset.iloc[2000:4000,:]
dataset = transfer_str_to_numeric_vals(sub_set_of_raw_dataset_2k_4k) 
X_train, X_test, y_train, y_test = split_to_train_and_test(dataset, category_col_name, test_ratio, rand_state)
best_max_dep, best_min_smpl_splt, best_f1_DT_params = find_best_decision_tree_params(X_train, y_train)
print("--- תוצאות GridSearchCV עבור Decision Tree ---")
print(f"עומק מקסימלי (max_depth) נבחר: {best_max_dep}")
print(f"פיצול מינימלי (min_samples_split) נבחר: {best_min_smpl_splt}")
print(f"ציון ה-F1 הממוצע הטוב ביותר: {best_f1_DT_params:.4f}")
```

Output:
```
--- תוצאות GridSearchCV עבור Decision Tree ---
עומק מקסימלי (max_depth) נבחר: 4
פיצול מינימלי (min_samples_split) נבחר: 5
ציון ה-F1 הממוצע הטוב ביותר: 0.9986
```

```python
# 3.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation, if used)
# === CODE TYPE: SELF TESTING
# ---- Add your additional tests here if needed:
###
### YOUR CODE HERE
###

```

<!-- grade_id: test3b-1_find_best_decision_tree_params, points: 0.1, grade: true, solution: false -->
```python
# 3.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 3.b. - Test 1 (name: test3b-1_find_best_decision_tree_params, points: 0.1) - Sanity")
print ("\t--->Testing the implementation of 'find_best_decision_tree_params' ...")

file_name = '.' + os.sep + 'data' + os.sep + 'mushrooms.csv'
test_ratio, rand_state = 0.2, 42
category_col_name = 'class'

try:
    raw_dataset = load_dataset(file_name)
    sub_set_of_raw_dataset_2k_4k = raw_dataset.iloc[2000:4000,:]
    dataset = transfer_str_to_numeric_vals(sub_set_of_raw_dataset_2k_4k) 
    X_train, X_test, y_train, y_test = split_to_train_and_test(dataset, category_col_name, test_ratio, rand_state)
    best_max_dep, best_min_smpl_splt, best_f1_DT_params = find_best_decision_tree_params(X_train, y_train)
except Exception as e:
    print ('You probably have a syntax error, we got the following exception:')
    print ('\tError Message:', str(e))
    print ('Try fixing your implementation')
    raise 
         
print ("\nGood Job!\nYou've passed the 1st test for the 'find_best_decision_tree_params' method :-)")
```

Output:
```
Part 3.b. - Test 1 (name: test3b-1_find_best_decision_tree_params, points: 0.1) - Sanity
	--->Testing the implementation of 'find_best_decision_tree_params' ...

Good Job!
You've passed the 1st test for the 'find_best_decision_tree_params' method :-)
```

<!-- grade_id: test3b-2_find_best_decision_tree_params, points: 0.4, grade: true, solution: false -->
```python
# 3.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 3.b. - Test 2 (name: test3b-2_find_best_decision_tree_params, points: 0.4)")
print ("\t--->Testing the implementation of 'find_best_decision_tree_params' ...")

file_name = '.' + os.sep + 'data' + os.sep + 'mushrooms.csv'
test_ratio, rand_state = 0.2, 42
category_col_name = 'class'

try:
    raw_dataset = load_dataset(file_name)
    sub_set_of_raw_dataset_2k_4k = raw_dataset.iloc[2000:4000,:]
    dataset = transfer_str_to_numeric_vals(sub_set_of_raw_dataset_2k_4k) 
    X_train, X_test, y_train, y_test = split_to_train_and_test(dataset, category_col_name, test_ratio, rand_state)
except Exception as e:
    print ('You probably have a syntax error, we got the following exception:')
    print ('\tError Message:', str(e))
    print ('Try fixing your implementation')
    raise 

found_assertion_combination = False
for i in range(50):
    try:
        best_max_dep, best_min_smpl_splt, best_f1_DT_params = find_best_decision_tree_params(X_train, y_train)
        ###
        ### AUTOGRADER TEST - DO NOT REMOVE
        ###
        assert best_max_dep==4 and best_min_smpl_splt==5,'Wrong best Decision-tree hyperparameter pair values'
        found_assertion_combination = True
        print ('... successful')
        break
    except Exception as e:
        print ('Try:',i+1,'failed')
        print ('\tError Message:', str(e))
        print ('Will try again ...')

assert found_assertion_combination,'Wrong best Decision-tree hyperparameter pair values'
         
print ("\nGood Job!\nYou've passed the 2nd test for the 'find_best_decision_tree_params' method :-)")
```

Output:
```
Part 3.b. - Test 2 (name: test3b-2_find_best_decision_tree_params, points: 0.4)
	--->Testing the implementation of 'find_best_decision_tree_params' ...
Try: 1 failed
	Error Message: Wrong best Decision-tree hyperparameter pair values
Will try again ...
... successful

Good Job!
You've passed the 2nd test for the 'find_best_decision_tree_params' method :-)
```

<!-- grade_id: test3b-3_find_best_decision_tree_params, points: 0.4, grade: true, solution: false -->
```python
# 3.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 3.b. - Test 3 (name: test3b-3_find_best_decision_tree_params, points: 0.4)")
print ("\t--->Testing the implementation of 'find_best_decision_tree_params' ...")
print ("\n\t====> Full grading test - the following test can not be seen before submission")
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

Output:
```
Part 3.b. - Test 3 (name: test3b-3_find_best_decision_tree_params, points: 0.4)
	--->Testing the implementation of 'find_best_decision_tree_params' ...

	====> Full grading test - the following test can not be seen before submission
```

### 3.c. Random forest - find best number of estimators parameter

### 3.c. Instructions
<u>method name</u>: <b>find_best_random_forest_num_estimators</b>
<pre>The following is expected:
--- Complete the 'find_best_random_forest_num_estimators' function to find best
    'n_estimators' hyperparameter values, 
     for the random-forest algorithm, running on the mushroom dataset's train-set.
     Again, use the grid search cross validation 'GridSearchCV' function.
Notes:
       You need to use the average f1 score, in order to choose the best  
           'n_estimators' hyperparameter for this dataset.
       Use 11,51,71 as possible values as possibilities for the 
           'n_estimators' parameter.
       You need to return the value of 'n_estimators' best value 
           and the best average f1 score.
       Use the 'make_scorer', as the function for 'scoring' parameter     
</pre>
<hr>
The return statement should look similar to the following statement:<br />
<b>return best_num_estimators, best_f1_val</b>

```python
# 3.c.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ASSISTANCE TO ANSWER 
# ---- Add assistance code here IF NEEDED:
###
### YOUR CODE HERE
###

```

```python
# 3.c.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ANSWER 

def find_best_random_forest_num_estimators(X_train, y_train):
    # 1. הגדרת רשת הפרמטרים עבור מספר העצים ביער
    param_grid = {'n_estimators': [11, 51, 71]}
    
    # 2. יצירת אובייקט ה-Scorer עבור מדד f1
    f1_scorer = make_scorer(f1_score)
    
    # 3. יצירת מודל יער אקראי בסיסי
    rf_clf = RandomForestClassifier()
    
    # 4. הגדרת ה-GridSearchCV (משתמש ב-5-Fold Cross Validation כברירת מחדל)
    grid_search = GridSearchCV(estimator=rf_clf, param_grid=param_grid, scoring=f1_scorer)
    
    # 5. הרצת החיפוש והאימון על נתוני האימון
    grid_search.fit(X_train, y_train)
    
    # 6. שליפת מספר העצים המנצח והציון הממוצע הכי גבוה
    best_num_estimators = grid_search.best_params_['n_estimators']
    best_f1_val = grid_search.best_score_
    
    return best_num_estimators, best_f1_val
```

```python
# 3.c.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: SELF TESTING
# Use the following code to test your implementation:
file_name = '.' + os.sep + 'data' + os.sep + 'mushrooms.csv'
test_ratio, rand_state = 0.2, 42
category_col_name = 'class'
raw_dataset = load_dataset(file_name)
sub_set_of_raw_dataset_2k_4k = raw_dataset.iloc[2000:4000,:]
dataset = transfer_str_to_numeric_vals(sub_set_of_raw_dataset_2k_4k) 
X_train, X_test, y_train, y_test = split_to_train_and_test(dataset, category_col_name, test_ratio, rand_state)
best_n_estimators, best_f1_RF_params = find_best_random_forest_num_estimators(X_train, y_train)
# --- add additional code to check your code if needed:
print("--- תוצאות GridSearchCV עבור Random Forest ---")
print(f"מספר העצים (n_estimators) הנבחר: {best_n_estimators}")
print(f"ציון ה-F1 הממוצע הטוב ביותר: {best_f1_RF_params:.4f}")
```

Output:
```
--- תוצאות GridSearchCV עבור Random Forest ---
מספר העצים (n_estimators) הנבחר: 51
ציון ה-F1 הממוצע הטוב ביותר: 0.9972
```

```python
# 3.c.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation, if used)
# === CODE TYPE: SELF TESTING
# ---- Add your additional tests here if needed:
###
### YOUR CODE HERE
###

```

<!-- grade_id: test3c-1_find_best_random_forest_num_estimators, points: 0.1, grade: true, solution: false -->
```python
# 3.c.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 3.c. - Test 1 (name: test3c-1_find_best_random_forest_num_estimators, points: 0.1) - Sanity")
print ("\t--->Testing the implementation of 'find_best_random_forest_num_estimators' ...")

file_name = '.' + os.sep + 'data' + os.sep + 'mushrooms.csv'
test_ratio, rand_state = 0.2, 42
category_col_name = 'class'

try:
    raw_dataset = load_dataset(file_name)
    sub_set_of_raw_dataset_2k_4k = raw_dataset.iloc[2000:4000,:]
    dataset = transfer_str_to_numeric_vals(sub_set_of_raw_dataset_2k_4k) 
    X_train, X_test, y_train, y_test = split_to_train_and_test(dataset, category_col_name, test_ratio, rand_state)
    best_n_estimators, best_f1_RF_params = find_best_random_forest_num_estimators(X_train, y_train)
except Exception as e:
    print ('You probably have a syntax error, we got the following exception:')
    print ('\tError Message:', str(e))
    print ('Try fixing your implementation')
    raise 
    
print ("Good Job!\nYou've passed the 1st test for the 'find_best_random_forest_num_estimators' function implementation :-)")
```

Output:
```
Part 3.c. - Test 1 (name: test3c-1_find_best_random_forest_num_estimators, points: 0.1) - Sanity
	--->Testing the implementation of 'find_best_random_forest_num_estimators' ...
Good Job!
You've passed the 1st test for the 'find_best_random_forest_num_estimators' function implementation :-)
```

<!-- grade_id: test3c-2_find_best_random_forest_num_estimators, points: 0.4, grade: true, solution: false -->
```python
# 3.c.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 3.c. - Test 2 (name: test3c-2_find_best_random_forest_num_estimators, points: 0.4)")
print ("\t--->Testing the implementation of 'find_best_random_forest_num_estimators' ...")

file_name = '.' + os.sep + 'data' + os.sep + 'mushrooms.csv'
test_ratio, rand_state = 0.2, 42
category_col_name = 'class'

try:
    raw_dataset = load_dataset(file_name)
    sub_set_of_raw_dataset_2k_4k = raw_dataset.iloc[2000:4000,:]
    dataset = transfer_str_to_numeric_vals(sub_set_of_raw_dataset_2k_4k) 
    X_train, X_test, y_train, y_test = split_to_train_and_test(dataset, category_col_name, test_ratio, rand_state)
except Exception as e:
    print ('You probably have a syntax error, we got the following exception:')
    print ('\tError Message:', str(e))
    print ('Try fixing your implementation')
    raise 

found_assertion_combination = False
for i in range(50):
    try:
        best_n_estimators, best_f1_RF_params = find_best_random_forest_num_estimators(X_train, y_train)
        ###
        ### AUTOGRADER TEST - DO NOT REMOVE
        ###
        assert best_n_estimators==71,'Wrong best num-estimator hyperparameter value for Random-Forest'
        found_assertion_combination = True
        print ('... successful')
        break
    except Exception as e:
        print ('\tError Message:', str(e))
        print ('Will try again, num-try:',i+1)

assert found_assertion_combination,'Wrong best num-estimator hyperparameter value for Random-Forest'
         
print ("Good Job!\nYou've passed the 2nd test for the 'find_best_random_forest_num_estimators' function implementation :-)")
```

Output:
```
Part 3.c. - Test 2 (name: test3c-2_find_best_random_forest_num_estimators, points: 0.4)
	--->Testing the implementation of 'find_best_random_forest_num_estimators' ...
	Error Message: Wrong best num-estimator hyperparameter value for Random-Forest
Will try again, num-try: 1
	Error Message: Wrong best num-estimator hyperparameter value for Random-Forest
Will try again, num-try: 2
... successful
Good Job!
You've passed the 2nd test for the 'find_best_random_forest_num_estimators' function implementation :-)
```

<!-- grade_id: test3c-3_find_best_random_forest_num_estimators, points: 0.4, grade: true, solution: false -->
```python
# 3.c.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 3.c. - Test 3 (name: test3c-3_find_best_random_forest_num_estimators, points: 0.4)")
print ("\t--->Testing the implementation of 'find_best_random_forest_num_estimators' ...")
print ("\n\t====> Full grading test - the following test can not be seen before submission")
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

Output:
```
Part 3.c. - Test 3 (name: test3c-3_find_best_random_forest_num_estimators, points: 0.4)
	--->Testing the implementation of 'find_best_random_forest_num_estimators' ...

	====> Full grading test - the following test can not be seen before submission
```

## 4. Select The best model
In this section you need to select the best model using the 
  average cross validation recall score.

### 4. Instructions
<u>method name</u>: <b>find_best_model</b>
<pre>The following is expected:
--- Complete the 'find_best_model' function to find best
    trained model out of the following three models:
    Model 1: Decision Trees model (with given 'max_depth' and 'min_samples_split' parameters)
    Model 2: (Gaussian) Naive Bayes model (with default parameters). 
    Model 3: SVM model (SVC with default parameters).
Notes:
       You need to use the average recall score, in order to choose the best  
           trained classification model.
       You need to return the best trained classifier object and the best average recall score.
          The average score, should be performed on a 10-fold cross validation.
          Use the 'cross_val_score' method to calculate the best average score.
</pre>
<hr>
The return statement should look similar to the following statement:<br />
<b>return best_clf, best_recall_val</b>

```python
# 4.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ASSISTANCE TO ANSWER 
# ---- Add assistance code here IF NEEDED:
###
### YOUR CODE HERE
###

```

```python
# 4.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ANSWER 

def find_best_model(X_train, y_train, max_depth_val, min_samples_split_val):
   # 1. הגדרת שלושת אובייקטי המסווגים המבוקשים עם השמות הנכונים של הפרמטרים
    dt_clf = DecisionTreeClassifier(max_depth=max_depth_val, min_samples_split=min_samples_split_val)
    nb_clf = GaussianNB()
    svm_clf = SVC()
    
    # 2. הרצת 10-Fold Cross Validation לכל מודל בעזרת מדד recall
    score_dt = cross_val_score(dt_clf, X_train, y_train, cv=10, scoring='recall').mean()
    score_nb = cross_val_score(nb_clf, X_train, y_train, cv=10, scoring='recall').mean()
    score_svm = cross_val_score(svm_clf, X_train, y_train, cv=10, scoring='recall').mean()
    
    # 3. מציאת הציון הגבוה ביותר והמודל המתאים לו
    best_recall_val = score_dt
    best_clf = dt_clf
    
    if score_nb > best_recall_val:
        best_recall_val = score_nb
        best_clf = nb_clf
        
    if score_svm > best_recall_val:
        best_recall_val = score_svm
        best_clf = svm_clf
        
    # 4. אימון המודל המנצח על כל נתוני האימון לפני ההחזרה
    best_clf.fit(X_train, y_train)
    
    return best_clf, best_recall_val
```

```python
# 4.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: SELF TESTING
# Use the following code to test your implementation:
file_name = '.' + os.sep + 'data' + os.sep + 'mushrooms.csv'
test_ratio, rand_state = 0.2, 42
category_col_name = 'class'
max_dep=4
min_smpl_splt=5
raw_dataset = load_dataset(file_name)
dataset = transfer_str_to_numeric_vals(raw_dataset) 
X_train, X_test, y_train, y_test = split_to_train_and_test(dataset, category_col_name, test_ratio, rand_state)
best_clf, best_recall_val=find_best_model(X_train, y_train, max_dep, min_smpl_splt)
# --- add additional code to check your code if needed:
print("--- תוצאות בחירת המודל המנצח (10-Fold Cross Validation) ---")
print(f"המודל שנבחר כהכי טוב: {type(best_clf).__name__}")
print(f"ציון ה-Recall הממוצע שלו: {best_recall_val:.4f}")

# וידוא שהמודל אכן עבר fit ומסוגל לנבא
try:
    test_preds = best_clf.predict(X_test[:5])
    print(f"בדיקת תקינות ניבוי (5 דגימות ראשונות): {test_preds}")
except Exception as e:
    print(f"שגיאה: המודל לא אומן כראוי. פירוט: {e}")
```

Output:
```
--- תוצאות בחירת המודל המנצח (10-Fold Cross Validation) ---
המודל שנבחר כהכי טוב: DecisionTreeClassifier
ציון ה-Recall הממוצע שלו: 0.9875
בדיקת תקינות ניבוי (5 דגימות ראשונות): [0 1 1 0 1]
```

```python
# 4.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation, if used)
# === CODE TYPE: SELF TESTING
# ---- Add your additional tests here if needed:
###
### YOUR CODE HERE
###

```

<!-- grade_id: test4-1_find_best_model, points: 0.1, grade: true, solution: false -->
```python
# 4.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 4. - Test 1 (name: test4-1_find_best_model, points: 0.1) - sanity")
print ("\t--->Testing the implementation of 'find_best_model' ...")

file_name = '.' + os.sep + 'data' + os.sep + 'mushrooms.csv'
test_ratio, rand_state = 0.2, 42
category_col_name = 'class'
max_dep=4
min_smpl_splt=5

try:
    raw_dataset = load_dataset(file_name)
    dataset = transfer_str_to_numeric_vals(raw_dataset) 
    X_train, X_test, y_train, y_test = split_to_train_and_test(dataset, category_col_name, test_ratio, rand_state)
    best_clf, best_recall_val=find_best_model(X_train, y_train, max_dep, min_smpl_splt)
except Exception as e:
    print ('You probably have a syntax error, we got the following exception:')
    print ('\tError Message:', str(e))
    print ('Try fixing your implementation')
    raise 

print ("Good Job!\nYou've passed the 1st test for the 'find_best_model' function implementation :-)")
```

Output:
```
Part 4. - Test 1 (name: test4-1_find_best_model, points: 0.1) - sanity
	--->Testing the implementation of 'find_best_model' ...
Good Job!
You've passed the 1st test for the 'find_best_model' function implementation :-)
```

<!-- grade_id: test4-2_find_best_model, points: 0.5, grade: true, solution: false -->
```python
# 4.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 4. - Test 2 (name: test4-2_find_best_model, points: 0.5)")
print ("\t--->Testing the implementation of 'find_best_model' ...")

file_name = '.' + os.sep + 'data' + os.sep + 'mushrooms.csv'
test_ratio, rand_state = 0.2, 42
category_col_name = 'class'
max_dep=4
min_smpl_splt=5

try:
    raw_dataset = load_dataset(file_name)
    dataset = transfer_str_to_numeric_vals(raw_dataset) 
    X_train, X_test, y_train, y_test = split_to_train_and_test(dataset, category_col_name, test_ratio, rand_state)
    best_clf, best_recall_val=find_best_model(X_train, y_train, max_dep, min_smpl_splt)
except Exception as e:
    print ('You probably have a syntax error, we got the following exception:')
    print ('\tError Message:', str(e))
    print ('Try fixing your implementation')
    raise 

###
### AUTOGRADER TEST - DO NOT REMOVE
###

assert type(best_clf) == sklearn.tree._classes.DecisionTreeClassifier , 'Wrong type for best classifier'

print ("Good Job!\nYou've passed the 2nd test for the 'find_best_model' function implementation :-)")
```

Output:
```
Part 4. - Test 2 (name: test4-2_find_best_model, points: 0.5)
	--->Testing the implementation of 'find_best_model' ...
Good Job!
You've passed the 2nd test for the 'find_best_model' function implementation :-)
```

<!-- grade_id: test4-3_find_best_model, points: 0.4, grade: true, solution: false -->
```python
# 4.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 4. - Test 3 (name: test4-3_find_best_model, points: 0.4)")
print ("\t--->Testing the implementation of 'find_best_model' ...")
print ("\n\t====> Full grading test - the following test can not be seen before submission")
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

Output:
```
Part 4. - Test 3 (name: test4-3_find_best_model, points: 0.4)
	--->Testing the implementation of 'find_best_model' ...

	====> Full grading test - the following test can not be seen before submission
```